```python
# Initialize Otter
import otter
grader = otter.Notebook("partA.ipynb")
```

# Programming
## Overview
In this file you will create some static visualizations AND an interactive dashboard
It is important to follow the instructions as this file will MOSTLY be autograded.
- PART 1 - Standard Static Vis
- PART 2 - Standard Static Vis
- PART 3 - Novel Visualization
- PART 4 - Interactive Dashboard

```python
import altair as alt
import pandas as pd
alt.renderers.enable('default')
```

## PART 1 - Standard Static Viz - Bar Chart

```python
### Read in the Data
data1 = pd.read_csv('https://raw.githubusercontent.com/kemiolamudzengi/2023W1-exam/main/part1.csv')
data1.head()
```


<div class="alert alert-info" style="color:black">

- Using the dataframe `data1`, Create the bar chart as depicted, we have provided some specs you might miss
- Using the `bar` mark
    - with a size of 16,
    - grey stroke,
    - strokeWidth of 1 and
    - cornerRadiusEnd of 8
- The height of each viz is 300, the width is 650
- `x` channel encodes 'State', `y` channel encodes 'Capacity', `color` channel encodes 'Region'
- The color scheme is set3
- Size of the viz title must be 24
- Legend title should be renamed to Region of Country
- For the x channel tick marks must be removed, as well as grid lines, the label should be oriented as shown
- The bars is sorted by the `Swept_Area` attribute.
- Name of the chart must be `bar_winds`

</div>

_Points:_ 10

```python

bar_winds = ...

# Show plot
bar_winds
```

```python
grader.check("q_part1")
```

## PART 2 - Standard Static Viz - Faceted Density Plots

```python
### Read in the data
data2 = pd.read_csv('https://raw.githubusercontent.com/kemiolamudzengi/2023W1-exam/main/part2.csv')
cut_to_keep = ['Fair', 'Ideal', 'Premium', 'Good']
data2 = data2.query("cut == @cut_to_keep")
data2.head()
```

<div class="alert alert-info" style="color:black">

- Using the dataframe `data2` create the faceted density plots shown, follow the following steps in this order below:
- Use transform_filter to only keep items that are of the color `J`, `D`, or `G`
- Create a density plot for the data, with the quantitative data being `carat`,
- Group it by the `cut` and overlay on top of each other
- Make sure the x scale is between 0 and 3
- With an opacity of 40%
- width is 200 pixels, height is default
- Use the Column (NOT FACET OR ROW) and group based on the `color` of the diamond
- Adhere to the chart styling shown in the image.
 </div>

_Points:_ 15

```python
dense_diamonds = ...


# Show dense_diamonds
dense_diamonds
```

```python
grader.check("q_part2")
```

## PART 3 - Novel Visualization - Altair Fan

### Overview
For this question, you have been given sample code (see cell below) AND documentation from Altair's API (see pdf)
The sample code is a starting point and provides you with a basis of which mark is being used.

Your goal is to create the visualization called Altair Fan
Don't panic, while you have never created this before, over the course of the term you have learned how to
- create data frames
- describe vizzes
- read the API and 
- create visualizations. 

In this question you get to demonstrate that knowledge.

```python
# SAMPLE CODE
# DO NOT EDIT THIS CELL
# LEAVE AS IS
# FEEL FREE TO COPY THE CODE INTO THE CELL BELOW AND THEN EDIT THERE

source = pd.DataFrame({"values": [12, 23, 47, 6, 52, 19]})

base = alt.Chart(source).encode(
    alt.Theta("values:Q").stack(True),
    alt.Radius("values").scale(type="sqrt", zero=True, rangeMin=20),
    color="values:N",
)

c1 = base.mark_arc(innerRadius=20, stroke="#fff")

c2 = base.mark_text(radiusOffset=10).encode(text="values:Q")

c1 + c2
```

<div class="alert alert-info" style="color:black">

- Determine what data is needed
- Determine which channels for `mark_arc` need to be encoded
- Determine (by reading the documentation) which attributes need to be set. 
- The width and height are 300 pixels. 
    - Use Pandas to create a dataframe, your column headings must be
    -`values` and `quad`
- Your dataframe must be called `data3`
- The color scheme is `set2`
- Your viz must be called `fan`
- HINT - each portion of the fan has 20 slivers (i.e., arcs)
    
 </div>

_Points:_ 25

```python
# create your dataframe here 
quad = ...
values = ...

data3 = ...
data3

```

```python
# Create your chart here
fan = ...

# Show Altair Fan
fan

```

```python
grader.check("q_novel")
```

## PART 4 - Interactive Dashboard

## Overview
This question is very similar to assquiab part 2 in structure.
You will create 3 visualizations, add interaction and place in a dashboard.
In this question you demonstrate your understanding of
- creating standard visualizations
- adding interaction to visualizations
- assemblying a dashboard

All the viz in this section use the data stored in `data4`

```python
data4 = pd.read_csv('https://raw.githubusercontent.com/kemiolamudzengi/2023W1-exam/main/part4.csv',parse_dates=['Date'])
data4.head()
```

## PART 4 - Question 1
<div class="alert alert-info" style="color:black">

Your task is to create a scatter chart.
The mark_circle should have a default size of 50.
The channels should be as follows
- y is `Mean_Temp` a quantitative attribute, change the title as shown
- x is `Total_Rain`, a quantitative attribute, change the title as shown
- color  is `Wind`, use the `tableau10` scheme
- The name of the chart is scatter_chart
- The size should be width = 300, height = 250
</div>

_Points:_ 5

```python
scatter_chart = ...


# Show the chart
scatter_chart

```

```python
grader.check("q4_part1")
```

## PART 4 - Question 2

<div class="alert alert-info" style="color:black">

Your task is to create a heatmap, using mark_rect
The encodings should be as follows
- x is `Arb` a quantitative attribute, change the title as shown
- y is the month of the Date, change the title as shown
- color  is the <b>mean</b> of the `Variation`, use the `greys` scheme
- The name of the chart is heatmap
- The size should be width = 300, height = 250

</div>

_Points:_ 5

```python
heatmap = ...

# Show heatmap

heatmap

```

```python
grader.check("q4_part2")
```

## PART 4 - Question 3

<div class="alert alert-info" style="color:black">

Your task is to create an area chart with the following encoding:
- y is `Cloudy` a quantitative attribute, change the title as shown
- x is `Date`, remove the axis labels and set tick to false
- chart properties width is 750 and height is 70
- tooltip with Date
- the name of the chart is `temporal_base`
- The mark should have `purple` as the fill with an opacity of 0.3
- Name of the area chart must be temporal_base

</div>

_Points:_ 10

```python

temporal_base = ...

# Show the chart
temporal_base

```

```python
grader.check("q4_part3")
```

## PART 4 Question 4: Interaction and Dashboard

<div class="alert alert-info" style="color:black">


1. Create a brush interval, named `time_brush`, constrained to the x, and with empty set to `True`,
2. Create a chart called `base` from the temporal_base. Add the `time_brush` as a parameter
3. Create a chart called `selected` from base, add `time_brush` using it as a <b>filter</b>, and make sure to have the mark_area have the color of 'purple'.
4. Concatenate the `base` and `selected` chart in that order (`base` first, `selected` second) and store in the variable `top`
5. Add the time_brush as the filter mechanism for the heatmap
6. Add the time_brush as a parameter for the scatter_chart, let the brush serve as `selector` for which circles have a default color of `lightgrey` if not selected, and the Wind color otherwise
7. Concatenate the two and store in variable `bottom`
8. Concatenate top and bottom and store in variable `dash`, visualize dash

</div>


_Points:_ 30

```python
time_brush = ...
base = ...
selected = ...
top = ...

bottom = ...

dash = ...

# SHOW THE DASHBOARD
dash
```

```python
grader.check("q4_part4")
```

<div class="alert alert-danger" style="color:black">

**Restart and run all before submitting**

Before submitting your notebook.
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see your submission on Gradescope.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for DSCI 320,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

Tests for q4_part1 (points: 5)
```
>>> assert scatter_chart is not None
>>> assert isinstance(scatter_chart, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q4_part2 (points: 5)
```
>>> assert heatmap is not None
>>> assert isinstance(heatmap, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q4_part3 (points: 10)
```
>>> assert temporal_base is not None
>>> assert isinstance(temporal_base, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q4_part4 (points: 30)
```
>>> assert time_brush is not None
>>> assert base is not None
>>> assert selected is not None
>>> assert top is not None
>>> assert bottom is not None
>>> assert dash is not None
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q_novel (points: 25)
```
>>> assert fan is not None
>>> assert data3 is not None
>>> assert isinstance(fan, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q_part1 (points: 10)
```
>>> assert bar_winds is not None
>>> assert data1 is not None
>>> assert isinstance(bar_winds, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q_part2 (points: 15)
```
>>> assert dense_diamonds is not None
>>> assert data2 is not None
>>> assert isinstance(dense_diamonds, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type
